<a href="https://colab.research.google.com/github/Samuel-Egbert31415/DS-Unit-1-Sprint-1-Data-Wrangling-and-Storytelling/blob/master/LS_DS_133_Multiple_Regression_Assignment_Samuel_Egbert_AG.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Lambda School Data Science - Unit 1 Sprint 3 Module 3

## Module Project: Multiple Linear Regression

### Learning Objectives

* model the relationship of multiple predictor variables to a single outcome
* conduct a t-test to determine the significance of individual variables in the model
* compare model fit using R-squared

## Part A: Introduction

### Statistical significance between head size and brain weight in healthy adult humans

The `Brainhead.csv` dataset provides information on 237 individuals who were subject to post-mortem examination at the Middlesex Hospital in London around the turn of the 20th century. Study authors used cadavers to see if a relationship between brain weight and other more easily measured physiological characterizes such as age, sex, and head size could be determined. The end goal was to develop a way to estimate a person’s brain size while they were still alive (as the living aren’t keen on having their brains taken out and weighed).

**We wish to determine if we can improve on our model of the linear relationship between head size and brain weight in healthy human adults.**

Source: R.J. Gladstone (1905). "A Study of the Relations of the Brain to the Size of the Head", Biometrika, Vol. 4, pp105-123.

### Use the above information to complete the following tasks.

**Task 1** - Load the data

Let's load the data! The URL has been provided as well as the imports for pandas and numpy.

* load your CSV file into a DataFrame named `df`

<!-- grade_id: cell-0434f06199d1ad11, points: 0, grade: false, solution: true -->
```python
# Task 1

import pandas as pd
import numpy as np

data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Brainhead/Brainhead.csv'

df = pd.read_csv(data_url)

# Print out your DataFrame
df.head()
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-1-ac2f92478458> in <module>()
      6 data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Brainhead/Brainhead.csv'
      7 
----> 8 df = pd.read_cvs(data_url)
      9 
     10 # Print out your DataFrame

AttributeError: module 'pandas' has no attribute 'read_cvs'
```

<!-- grade_id: cell-1c36b133f6ac028d, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(df, pd.DataFrame), 'Have you created a DataFrame named `df`?'
assert len(df) == 237

```

**Task 2** - Fit OLS MOdel

Run the OLS model relating head size to brain weight - you should have done this in your project yesterday. The `statsmodels.formula.api` import has been provided. After you fit your model, print out your model results.

* Fit an OLS model and name your variable `model`.

*Hint: Make sure to use the format Y ~ X*

<!-- grade_id: cell-2aa2fa6eba350a2c, points: 0, grade: false, solution: true -->
```python
# Task 2

# Import statsmodels - DON'T Delete
from statsmodels.formula.api import ols

# Model input format
# format Y ~ X

# YOUR CODE HERE
raise NotImplementedError()

# Print the model summary
print(model.summary())
```

<!-- grade_id: cell-bb49143b30a4fcee, points: 1, grade: true, solution: false -->
```python
# Task 2 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 3** - Adjusted $R^2$

Using your model summary, what is the adjusted $R^2$ for this model?  

* Assign the value of adjusted $R^2$ to the variable `adj_R2` (use three decimal places, e.g. 0.123). Your variable should be a float.

<!-- grade_id: cell-0b8d856cc141c2a7, points: 0, grade: false, solution: true -->
```python
# Task 3

# YOUR CODE HERE
raise NotImplementedError()

# Print the adjusted R^2 value
print('The adjusted R^2 is: ', adj_R2)
```

<!-- grade_id: cell-70088d009573d8ab, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 3** - Short answer

1. Explain what this value means in a complete sentence.

---

1. YOUR ANSWER

**Task 4** - Encoding new variables

From the data dictionary for this data set:

Gender
* 1 = male
* 2 = female

Age
* 1 = younger than age 46
* 2 = age 46 and older

Right now age and gender are coded with 1 and 2.  Just to be safe when we run our model, we want to subtract "1" from Age and Gender so that "1" is now coded as "0" and "2" is now coded as "1".  Create two new columns in the DataFrame to do this - don't overwrite the original columns!

**Your tasks:**

* Create a new Series (column) for age and give it the column name `age_01`.
* Create a new Series (column) for gender and give it the column `gender_01`.

<!-- grade_id: cell-ad1f0446a3bc99b1, points: 0, grade: false, solution: true -->
```python
# Task 4

# YOUR CODE HERE
raise NotImplementedError()

# View the DataFrame
df.head()
```

<!-- grade_id: cell-84c1602875b4b2fd, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test

assert df.shape == (237, 6), 'Did you add two new columns (and only two)?'

```

**Task 5** - Multiple Regression (adding in "Age")

We're going to try to improve our model by fitting a model on more than one independent parameter.

* Fit a model to Y (brain weight) and X1 and X2 (head size and age group); assign your model to the variable `model_2` (make sure to use `Age_01` for the age variable)

*Hint: Enter the model in the format Y ~ X1 + X2*

<!-- grade_id: cell-f10e617763776ff5, points: 0, grade: false, solution: true -->
```python
# Task 5

# Model input format: Y ~ X1 + X2*

# YOUR CODE HERE
raise NotImplementedError()

# Print the model summary
print(model_2.summary())
```

<!-- grade_id: cell-ba2bb42fefe49970, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 6** - Statistical significance

* Find the p-value for `Age_01` in your model and assign it to the variable `pval_age` (use three decimal places, e.g. 0.123). Your value should be a float.

At the alpha = 0.05 level, is age group statistically significantly associated with brain weight after accounting for head size?  Explain your answer in at least one complete sentence.

ANSWER:

At the alpha = 0.01 level, is age group statistically significantly associated with brain weight after accounting for head size? Explain your answer in at least one complete sentence.

ANSWER:


<!-- grade_id: cell-e76b59ebd9a33e39, points: 0, grade: false, solution: true -->
```python
# Task 6

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-45d090dbd7235338, points: 1, grade: true, solution: false -->
```python
# Task 6 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 7** - Adjusted $R^2$

What is the adjusted $R^2$ for the model relating head size and age group to brain weight?  

* Assign the adjusted $R^2$ value to `adj_R2_new`. This value should be a float.

Does adding age group to the model increase the amount of variability in brain weight explained by the model?

ANSWER:

<!-- grade_id: cell-39e56b2549fd6715, points: 0, grade: false, solution: true -->
```python
# Task 7

# YOUR CODE HERE
raise NotImplementedError()

# Print the adjusted R^2 new value
print('The adjusted R^2 new is: ', adj_R2_new)
```

<!-- grade_id: cell-1cce579f5ceec835, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

# Hidden tests - you will see the results when you submit to Canvas
```

**Task 8** - Multiple Regression (adding in Gender)

We're going to try another multiple regression to improve our model!

* Fit a model to Y (brain weight) and X1 and X2 (head size and gender); assign your model to the variable `model_3` (make sure to use `Gender_01` for the age variable).

*Hint: Enter the model in the format Y ~ X1 + X2*

<!-- grade_id: cell-6e2d7e8b615b59d7, points: 0, grade: false, solution: true -->
```python
# Task 8

# Model input format: Y ~ X1 + X2*

# YOUR CODE HERE
raise NotImplementedError()

# Print the model summary
print(model_3.summary())
```

<!-- grade_id: cell-5e7bef722b7fe2ec, points: 1, grade: true, solution: false -->
```python
# Task 8 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 9** - Significance (short answer)

At the alpha = 0.05 level, is gender statistically significantly associated with brain weight after accounting for head size?  Explain your answer in at least one complete sentence.

ANSWER:

**Task 10** - Linear model equation (short answer)

Write the linear model (insert the coefficients in the correct places) used to model the relationship between head size, age group and predicted brain weight. *Make sure to use the correct model!*

EQUATION: 

**Task 11** - Create a plot!

Plot the relationship between head size and brain weight on a scatterplot. Use different color dots to represent older and younger individuals. You can use `matplotlib` or `seaborn`.

<!-- grade_id: cell-5fe9394388a9edfb, points: 0, grade: false, solution: true -->
```python
# Task 11

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 12** - Make predictions

Estimate the brain weight of a 60-year-old with a head size of 4000 cubic cm. **Make sure to select the correct model you fit above!**

* assign the head size to the variable `x1` (integer)
* assign the age group (either a 0 or 1) to the variable `x2` (integer)
* assign the intercept to `intercept` (float)
* assign the slope parameter for the head size to `slope1` (float)
* assign the slope parameter for age to `slope2` (float)
* assign the prediction to the variable `predicted` (float)

<!-- grade_id: cell-c6a58b3a206a56d1, points: 0, grade: false, solution: true -->
```python
# Task 12

# YOUR CODE HERE
raise NotImplementedError()

# Print out variables to troubleshoot!
# print('slope 1:', slope1)

# View your prediction
print('The predicted value for brain size is:', predicted)
```

<!-- grade_id: cell-4d1b9dcc5f82390f, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 13** - Make another prediction

Estimate the brain weight of a 20-year-old with a head size of 3900 cubic cm. Use the same slope and intercept parameters you used above. But - change your `x1` and `x2` values!

* change the values for `x1` and `x2` (both are integers)
* assign your prediction to the variable `predicted_2` (float)

<!-- grade_id: cell-0f16313e839553e8, points: 0, grade: false, solution: true -->
```python
# Task 13

# YOUR CODE HERE
raise NotImplementedError()

# View your prediction
print('The predicted value for brain size is:', predicted_2)

```

<!-- grade_id: cell-82ff9c0933833051, points: 1, grade: true, solution: false -->
```python
# Task 13 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

## Part B

### Sleep Data

Use the following information to answer the rest of the tasks in this project.

Researchers recorded data on sleep duration as well as a set of ecological and constitutional variables  for a selection of mammal species. This data appears in the Sleep.csv dataset. 

Source: Allison, T. and Cicchetti, D. (1976), "Sleep in Mammals: Ecological and Constitutional Correlates",  Science, November 12, vol. 194, pp. 732-734.

**Data Dictionary:**

| Variable Name |            Description           |                                                                 Details                                                                 |              |   |
|:-------------:|:--------------------------------:|:---------------------------------------------------------------------------------------------------------------------------------------:|--------------|---|
| Animal        |            Animal name           |                                                            Character [string]                                                           |    |   |
| Body          |            Body weight           |                                                               Numeric [kg]                                                              |       |   |
| Brain         |           Brain weight           |                                                               Numeric [g]                                                               |  |   |
| SWsleep       | Slow wave (“non-dreaming”) sleep | Numeric [hours]                                                                                                                         |              |   |
| Parasleep     | Paradoxical (“dreaming”) sleep   | Numeric [hours]                                                                                                                         |              |   |
| Totsleep      | Total sleep                      | Numeric [hours]                                                                                                                         |              |   |
| Life          | Maximum life expectancy          | Numeric [years]                                                                                                                         |              |   |
| Gest          | Gestation time                   | Numeric [days]                                                                                                                          |              |   |
| Pred          | Predation index                  | Numeric [1 – 5] 1 = least likely to be preyed upon, 5 = most likely to be preyed upon                                                   |              |   |
| Exposure      | Sleep exposure index             | Numeric [1 – 5] 1 = least amount of exposure during sleep (mammal sleeps indoors or in a den), 5 = most amount of exposure during sleep |              |   |
| Danger        | Overall danger index             | Numeric [ 1 – 5] 1 = least amount of danger from other animals, 5 = most amount of danger from other animals                            |              |   |



**Task 14**

Before we can look at the data, we need to load in the data. The URL has been provided.

* Load in the CSV file as a DataFrame and assign it to the variable `df_sleep`.
* Make sure to view the DataFrame!

<!-- grade_id: cell-3a230410ee200f46, points: 0, grade: false, solution: true -->
```python
# Task 14

# Dataset URL - don't delete!
data_url_2 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Sleep/Sleep.csv'

# YOUR CODE HERE
raise NotImplementedError()

# Print out your DataFrame
df_sleep.head()
```

At the end of the guided project, we'd decided that the best model of dreaming sleep has the two variables gestation time and danger index.  

**Task 15** - Plot

Plot the relationship of body size to dreaming sleep. Turn off the CI on the plot (`ci = None`). Describe the shape of the relationship.

YOUR ANSWER:

<!-- grade_id: cell-a4468f535024cc41, points: 0, grade: false, solution: true -->
```python
# Task 15

# YOUR CODE HERE
raise NotImplementedError()

```

**Task 16** - Transformation: take the log

Does it appear that a log transformation of the `Body` variable would improve the model fit? Let take the log and then look at the plot again.

* Create a new column with the log of `Body` and add it as a column to `df_sleep` with the name `log_body`.
* Create a new column with the log of gestational time and add it as a column to `df_sleep` with the name `log_gest` 

*Hint: use the natural log `np.log()`*

<!-- grade_id: cell-6a290216ec0e3aaa, points: 0, grade: false, solution: true -->
```python
# Task 16

# YOUR CODE HERE
raise NotImplementedError()

# Print out the updated DataFrame
df_sleep.head()
```

<!-- grade_id: cell-db6506b6a01383b4, points: 1, grade: true, solution: false -->
```python
# Task 16 - Test
assert df_sleep.shape == (42, 13), 'Make sure you added two columns.'

# Hidden tests - you will see the results when you submit to Canvas
```

<!-- grade_id: cell-87bfbdd705aa0078, points: 0, grade: false, solution: true -->
```python
# Task 16 - Plot
# NOT autograded

# YOUR CODE HERE
raise NotImplementedError()
```

**Task 17** - Multiple Regression Model

Run an OLS model using gestation time, danger index and the log of body size.  
(This task will not be autograded)

Is the log of body size statistically significantly associated with time spent in dreaming sleep after accounting for the other two factors?

YOUR ANSWER

<!-- grade_id: cell-8b8a4ccb9569bca1, points: 0, grade: false, solution: true -->
```python
# Task 17

# Enter the model in the format Y ~ X1 + X2 + X3

# YOUR CODE HERE
raise NotImplementedError()

#Print the model summary

```

**Task 18** - Plot

Plot the relationship of gestational time and dreaming sleep with different colors for animals who experience different amounts of sleep exposure.

(This task will not be autograded)

<!-- grade_id: cell-64c517c8f8149502, points: 0, grade: false, solution: true -->
```python
# Task 18
# NOT autograded

# YOUR CODE HERE
raise NotImplementedError()

plt.show()
```

**Task 19** - Multiple Regression Model

Run an OLS model using the log of gestation time and sleep exposure index to predict the amount of time spent in dreaming sleep.  Is sleep exposure index a better predictor of amount of time spent in dreaming sleep than danger index?

(This task will not be autograded)

ANSWER:

<!-- grade_id: cell-f3d18879a9ddf4d0, points: 0, grade: false, solution: true -->
```python
# Task 19
# NOT autograded

# YOUR CODE HERE
raise NotImplementedError()

# Print the model summary

```

**Task 20** - Plot

Plot the relationship of gestational time and dreaming sleep with different colors for animals who experience different amounts of sleep exposure.

(This task will not be autograded)

ANSWER: 

<!-- grade_id: cell-4c80b7b174155807, points: 0, grade: false, solution: true -->
```python
# Task 20
# NOT autograded

# YOUR CODE HERE
raise NotImplementedError()


plt.show()
```

**Task 21** - Multiple Regression Model

Run an OLS model using the log of gestation time and predation index to predict the amount of time spent in dreaming sleep.  Is predation index a better predictor of amount of time spent in dreaming sleep than danger index?

(This task will not be autograded)

ANSWER: 

<!-- grade_id: cell-1c00b428d5be8030, points: 0, grade: false, solution: true -->
```python
# Task 21
# NOT autograded

# YOUR CODE HERE
raise NotImplementedError()

# Print the model summary

```

## Portfolio Project Assignment

Get (and give) constructive feedback to a peer.  Swap projects with at least one other student, and give each other constructive feedback.  Make sure to tell them what you like about their project as well as any suggestions you have for improvement.