<a href="https://colab.research.google.com/github/Samuel-Egbert31415/DS-Unit-1-Sprint-1-Data-Wrangling-and-Storytelling/blob/master/LS_DS_121_ttests_Assignment_AG.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## Autograded Notebook (Canvas & CodeGrade)

This notebook will be automatically graded. It is designed to test your answers and award points for the correct answers. Following the instructions for each Task carefully.

### Instructions

* **Download this notebook** as you would any other ipynb file
* **Upload** to Google Colab or work locally (if you have that set-up)
* **Delete `raise NotImplementedError()`**
* Write your code in the `# YOUR CODE HERE` space
* **Execute** the Test cells that contain `assert` statements - these help you check your work (others contain hidden tests that will be checked when you submit through Canvas)
* **Save** your notebook when you are finished
* **Download** as a `ipynb` file (if working in Colab)
* **Upload** your complete notebook to Canvas (there will be additional instructions in Slack and/or Canvas)

# Lambda School Data Science - Unit 1 Sprint 2 Module 1

## Hypothesis Testing - One and two-sample t-Tests

### Objectives

* Explain the purpose of a t-test and identify applications
* Use a t-test for independence to test for a statistically significant association between two categorical variables
* Use a t-test p-value to draw the correct conclusion about the null and alternative hypothesis

### Introduction

Mosquito nets have traditionally been an important tool to prevent mosquito bites in parts of the world where malaria is endemic. However, it may not be practical for an army that is on the move to set up and carry mosquito nets each night and day. Impregnating soldiers’ uniforms with insect repellent solves the mobility problem but also has drawbacks. First, the insect repellent quickly becomes ineffective with repeated washing and ironing and must be frequently reapplied. Second, in hot and humid climates the insect repellent can be absorbed through the skin, and the long-term effects of this exposure are unknown. One compromise is to have soldiers apply patches treated with insect repellent to their clothing. These patches would last longer because they would not be washed or ironed and would not expose the entire body to the insect repellent.

### Dataset Description

The `Mosquito.xlsx` dataset contains data recorded in an experiment conducted on male soldiers in the Indian Army who were stationed in the Tezpur/Solmara garrison in Northeast India. Thirty soldiers were randomly selected to receive one of three types of mosquito single repellent patch. After giving informed consent, the study participants affixed the patches at predetermined points on their uniforms and research assistants (who were blinded to the type of repellent used) counted the number of times a mosquito landed on each individual in an hour. 

Medical officers with the Indian Army have recorded data on mosquito bites and related illness for many years and can say with authority that the mean number of mosquito touches for soldiers not wearing any mosquito repellent is 8.2 per hour.**We wish to determine if wearing a single repellent patch changes the mean number of mosquito touches for soldiers compared to not wearing any mosquito repellant.**

*Adapted from: A. Bhatnagar and V.K. Mehta (2007). "Efficacy of Deltamethrin and Cyfluthrin Impregnated Cloth Over Uniform Against Mosquito Bites," Medical Journal Armed Forces India, Vol. 63, pp. 120-122.*

**Task 1** - Load the data

Let's load the data! The URL has been provided as well as the imports for pandas and numpy.

* load your CSV file into a DataFrame named `df_mosquito`

<!-- grade_id: cell-dae3789e00dd9b29, points: 0, grade: false, solution: true -->
```python
# Task 1
import pandas as pd
import numpy as np

# URL for the dataset
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Mosquito/Mosquito.csv'

df_mosquito = pd.read_csv(data_url)


# Print out your DataFrame
df_mosquito.head()
```

Output:
```
ID  Mosq_count
0   1           4
1   2          10
2   3          13
3   4           0
4   5          11
```

**Task 1 Test**

<!-- grade_id: cell-a347d8dd7572e6d4, points: 1, grade: true, solution: false -->
```python
# Task 1 - Test

assert isinstance(df_mosquito, pd.DataFrame), 'Have you created a DataFrame named `df_mosquito`?'
assert len(df_mosquito) == 90

```

**Task 2 -** Statistical hypotheses

Write the null and alternative hypotheses using the experiment information described above.

This task will not be autograded - but it is part of completing the challenge.

**Task 2 ANSWER:**

FILL IN YOUR ANSWERS BELOW

Null hypothesis: There is no effect on the number of mosquitos landing on soldiers by having them wear a patch vs not having any repellent at all. The mean should not be different from 8.2 lands per hour

Alternative hypothesis: 

$H_0: \mu =$ 8.2 lands

$H_a: \mu \neq$ 8.2 lands

**Task 3** - Calculate the mean

* Calculate the mean number of mosquito touches in the sample. Assign your answer to the variable `mosquito_touch_mean`.

<!-- grade_id: cell-22ced5cd2cae6230, points: 0, grade: false, solution: true -->
```python
# Task 3

mosquito_touch_mean = df_mosquito['Mosq_count'].mean()
print(mosquito_touch_mean)
```

Output:
```
8.011111111111111
```

**Task 3 Test**

<!-- grade_id: cell-3970fbb1a3eaefb5, points: 1, grade: true, solution: false -->
```python
# Task 3 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 4** - Calculate the standard deviation

* Calculate the standard deviation of the number of mosquito touches in the sample. Assign your answer to `mosquito_touch_std`.

<!-- grade_id: cell-9c4fe58a9fe2d7aa, points: 0, grade: false, solution: true -->
```python
# Task 4

mosquito_touch_std = df_mosquito['Mosq_count'].std()
print(mosquito_touch_std)
```

Output:
```
3.2825532828777257
```

**Task 4 Test**

<!-- grade_id: cell-4939c4289c39b4e3, points: 1, grade: true, solution: false -->
```python
# Task 4 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 5** - One-sample t-test

* Conduct a 1-sample t-test to test your hypotheses. Assign your t-test result to the variable `mosquito_pval`.

*Hint: The `stats.ttest_1samp()` function returns two values; read the documentation to select the p-value*

<!-- grade_id: cell-44b3f2a39081c5df, points: 0, grade: false, solution: true -->
```python
# Task 5

# Use the 'ttest_1samp' from the stats package
from scipy import stats

mosquito_pval = stats.stats.ttest_1samp(df_mosquito['Mosq_count'], 8.2)
mosquito_pval = mosquito_pval[1]
print(mosquito_pval)
```

Output:
```
0.5864980356272131
```

**Task 5 Test**

<!-- grade_id: cell-3be2eb1f16e366e6, points: 1, grade: true, solution: false -->
```python
# Task 5 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 6**

Report your conclusion at the 0.05 significance level. Write your answer in the cell below.

This task will not be autograded - but it is part of the project!

**Task 6 ANSWER:**

There is no significant evidences given a p = 0.05 significance to reject the null hypotheses that the patches have no effect: mosquito_ppval = 0.58 >0.05

## Use the following information to complete Tasks 7-13

### Introduction

More than 14,000 people finished the 2020 Disney Marathon held on January 12. The results by age and gender group are included in the `Disney.csv dataset`. 

**We wish to determine if the mean finishing time for male and female marathon runners is the same or if there is a difference in the mean finishing time between male and female marathon runners.**


[Source: Track Shack. 2020 Disney Marathon Race Results](https://www.trackshackresults.com/disneysports/results/wdw/wdw20/mar_results.php)

**Task 7** - Load the next dataset

Let's load the data! The URL has been provided.

* load your CSV file into a DataFrame named `df_disney`

<!-- grade_id: cell-2b389fa658a51cd2, points: 0, grade: false, solution: true -->
```python
# Task 7

# URL for Disney marathon dataset
data_url2 = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_1/Disney_Marathon/Disney.csv'

df_disney = pd.read_csv(data_url2)
df_disney
```

Output:
```
ID gender  age   group      time
0          1      M   30  M30-34  2.371944
1          2      M   26  M25-29  2.450556
2          3      M   32  M30-34  2.457778
3          4      M   35  M35-39  2.655833
4          5      M   26  M25-29  2.736111
...      ...    ...  ...     ...       ...
14101  14102      F   39  F35-39  7.320278
14102  14103      F   54  F50-54  7.340556
14103  14104      M   39  M35-39  7.383333
14104  14105      M   52  M50-54  7.400000
14105  14106      M   46  M45-49  7.406944

[14106 rows x 5 columns]
```

**Task 7 Test**

<!-- grade_id: cell-c1f16b75c543e0d1, points: 1, grade: true, solution: false -->
```python
# Task 7 - Test

assert isinstance(df_disney, pd.DataFrame), 'Have you created a DataFrame named `df_disney`?'
assert len(df_disney) == 14106

```

**Task 8** - Write the hypothesis

Write the null and alternative hypotheses for the scenario described above in words and symbols.

This task will not be autograded - but it is part of completing the challenge.

**Task 8 ANSWER:**

FILL IN YOUR ANSWERS BELOW

Null hypothesis: The mean finish times between groups Male and Female are not different

Alternative hypothesis: the mean finish times between groups Male and Femal are different 

$H_0: \mu_Male = \mu_Female$

$H_a: \mu_Male \neq \mu_Female$

**Task 9** - Create new series from a DataFrame

Create two separate Series (a pandas DataFrame column is a Series):

* one containing finishing times for male participants (`male_finish`)
* one containing finishing times for female participants (`female_finish`)

<!-- grade_id: cell-4290d3dd2f668a25, points: 0, grade: false, solution: true -->
```python
# Task 9

male_finish = df_disney[df_disney['gender']=='M']
print(male_finish['gender'].value_counts())
print(male_finish)

female_finish = df_disney[df_disney['gender']=='F']
print(female_finish['gender'].value_counts())
print(female_finish)

```

Output:
```
M    6577
Name: gender, dtype: int64
          ID gender  age   group      time
0          1      M   30  M30-34  2.371944
1          2      M   26  M25-29  2.450556
2          3      M   32  M30-34  2.457778
3          4      M   35  M35-39  2.655833
4          5      M   26  M25-29  2.736111
...      ...    ...  ...     ...       ...
14092  14093      M   68    M65+  7.280556
14099  14100      M   46  M45-49  7.403333
14103  14104      M   39  M35-39  7.383333
14104  14105      M   52  M50-54  7.400000
14105  14106      M   46  M45-49  7.406944

[6577 rows x 5 columns]
F    7529
Name: gender, dtype: int64
          ID gender  age   group      time
16        17      F   37  F35-39  2.905278
27        28      F   22  F18-24  2.985278
61        62      F   44  F40-44  3.207778
69        70      F   31  F30-34  3.244444
85        86      F   28  F25-29  3.290833
...      ...    ...  ...     ...       ...
14097  14098      F   51  F50-54  7.380000
14098  14099      F   55  F55-59  7.38583
```

**Task 9 Test**

<!-- grade_id: cell-5384320ccd7f5048, points: 1, grade: true, solution: false -->
```python
# Task 9 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 10** - Calculate the mean finishing times

* Calculate the mean finishing time for male and female participants separately. Name your variables `male_finish_mean` and `female_finish_mean`.

<!-- grade_id: cell-f6c55ea4a2dba463, points: 0, grade: false, solution: true -->
```python
# Task 10

male_finish_mean = male_finish['time'].mean()
print(male_finish_mean)

female_finish_mean = female_finish['time'].mean()
print(female_finish_mean)


```

Output:
```
5.799159782400031
6.291841988756132
```

**Task 10 Test**

<!-- grade_id: cell-71bc6659db96de45, points: 1, grade: true, solution: false -->
```python
# Task 10 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 11** - Calculate the standard deviation

* Calculate standard deviation of the mean finishing time for male and female participants separately. Name your variables `male_finish_std` and `female_finish_std`.

<!-- grade_id: cell-3ac04d8339ec03c3, points: 0, grade: false, solution: true -->
```python
# Task 11

male_finish_std = male_finish['time'].std()
print(male_finish_std)

female_finish_std = female_finish['time'].std()
print(female_finish_std)
```

Output:
```
1.100676340530379
0.896690100351361
```

**Task 11 Test**

<!-- grade_id: cell-a2de7fe5b67a1886, points: 1, grade: true, solution: false -->
```python
# Task 11 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 12** - Conduct a 2-sample t-test

Conduct a 2-sample t-test to test your hypotheses:

* Assign the t-statistic to a variable called `disney_tval`
* Assign the p-value to a variable called `disney_pval`

Hint: The function returns two values and you can assign them with one line (example):

`variable1, variable2` = `some.function(arguments)`

<!-- grade_id: cell-d391e7e149994d87, points: 0, grade: false, solution: true -->
```python
# Task 12

disney_tval = stats.stats.ttest_ind(male_finish['time'], female_finish['time'])
disney_tval = disney_tval[0]
print(disney_tval)
```

Output:
```
-29.27857393997243
```

**Task 12 Test**

<!-- grade_id: cell-35a1e3947ac6388a, points: 1, grade: true, solution: false -->
```python
# Task 12 - Test
# Hidden tests - you will see the results when you submit to Canvas
```

**Task 13** - Report your result

Report your conclusion at the 0.05 significance level.

This task will not be autograded - but it is part of completing the challenge.

With a significance level of p = 0.05, we can reject the null hypothesis that the mean Male and Female finish times in the disney marathon are different

### In your own words, answer the following questions:

These tasks will not be autograded - but they are part of completing the challenge.

**Task 14:** Explain the Central Limit Theorem.

As you continue to run trials on a random event, the frequency of the event will eventually be equal to the probability of the event as trials go to infinity

**Task 15:** Describe the Normal Distribution.

The gaussian distribution is a distrubution that is inevitable given enough samples of the population.  It has a specific structure and size given the parent population and sample sizes

**Task 16:** Describe the relationship between the Normal distribution and the t-distribution.

t-distribution is the resulting distrubution from sampling means from a parent distribution, except we didn't sample to infinity

**Task 17:** Write about who William Sealy Gosset was.

Beer guy working for Guinness beer, and he developed the "student" t-test doing his research


## Portfolio Project Milestone

Not sure what to do yet, I will figure it out though 

You should have selected a dataset and have a good idea what your research question is by the end of the day.  If you don't - please ask for help from your instructor, track team or mentor.